# Numpy Exercise 1

## Imports

```python
import numpy as np
import vizarray as va
```

## Checkerboard

Write a Python function that creates a square `(size,size)` 2d Numpy array with the values `0.0` and `1.0`:

* Your function should work for both odd and even `size`.
* The `0,0` element should be `1.0`.
* The `dtype` should be `float`.

<!-- grade_id: numpyex01a, points: 0, grade: false, solution: true -->
```python
def checkerboard(size):
    """Return a 2d checkboard of 0.0 and 1.0 as a NumPy array"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: numpyex01b, points: 6, grade: true, solution: false -->
```python
a = checkerboard(4)
assert a[0,0]==1.0
assert a.sum()==8.0
assert a.dtype==np.dtype(float)
assert np.all(a[0,0:5:2]==1.0)
assert np.all(a[1,0:5:2]==0.0)

b = checkerboard(5)
assert b[0,0]==1.0
assert b.sum()==13.0
assert np.all(b.ravel()[0:26:2]==1.0)
assert np.all(b.ravel()[1:25:2]==0.0)
```

Output:
```
<ipython-input-3-4aa63174c227> in <module>()
----> 1 a = checkerboard(4)
      2 assert a[0,0]==1.0
      3 assert a.sum()==8.0
      4 assert a.dtype==np.dtype(float)
      5 assert np.all(a[0,0:5:2]==1.0)

<ipython-input-2-b94219dc221c> in checkerboard(size)
      2     """Return a 2d checkboard of 0.0 and 1.0 as a NumPy array"""
      3     # YOUR CODE HERE
----> 4     raise NotImplementedError()

NotImplementedError:
```

Use `vizarray` to visualize a checkerboard of `size=20` with a block size of `10px`.

<!-- grade_id: numpyex01c, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-4-a3eb14613b08> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

Use `vizarray` to visualize a checkerboard of `size=27` with a block size of `5px`.

<!-- grade_id: numpyex01d, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-5-a3eb14613b08> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

```python

```