# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: projecteuler52a, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    # gota be the same length
    lst_strX = list(str(x)) # ['h','e',...]
    lst_strY = list(str(y)) # ['h','e',...]
    if (len(lst_strX) != len(lst_strY)):
        return False
    for i in range(0,len(lst_strX)):
        if (lst_strX.count(lst_strX[i]) == lst_strY.count(lst_strX[i])):
            i += 1
            continue
        else:
            return False
    return True
```

```python
print (same_digits(1235,3215))
```

Output:
```
True
```

<!-- grade_id: projecteuler52b, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: projecteuler52c, points: 6, grade: true, solution: true -->
```python
i = 0
x = True
while (x == True):
    i += 1
    if (same_digits(i*2,i*3) and
        same_digits(i*2,i*4) and
        same_digits(i*2,i*5) and
        same_digits(i*2,i*6) ):
        print (i)
        break
    if (i > 1000000):
        x = False
```

Output:
```
142857
```

```python
i = 142857
i*3, i*2, i*4, i*5, i*6
```

Output:
```
(428571, 285714, 571428, 714285, 857142)
```

```python

```