# Complete graph Laplacian

## Imports

```python
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
```

In discrete mathematics a [Graph](http://en.wikipedia.org/wiki/Graph_%28mathematics%29) is a set of *vertices* or *nodes* that are connected to each other by *edges* or *lines*. If those *edges* don't have directionality, the graph is said to be *undirected*. Graphs are used to model social and communications networks (Twitter, Facebook, Internet) as well as natural systems such as molecules.

A [Complete Graph](http://en.wikipedia.org/wiki/Complete_graph), $K_n$ on $n$ nodes has an edge that connects each node to every other node.

Here is $K_5$:

```python
import networkx as nx
K_5=nx.complete_graph(5)
nx.draw(K_5)
```

The [Laplacian Matrix](http://en.wikipedia.org/wiki/Laplacian_matrix) is a matrix that is extremely important in graph theory and numerical analysis. It is defined as $L=D-A$. Where $D$ is the [degree matrix](https://en.wikipedia.org/wiki/Degree_matrix) and $A$ is the [adjecency matrix](https://en.wikipedia.org/wiki/Adjacency_matrix). For the purpose of this problem you don't need to understand the details of these matrices, although their definitions are relatively simple (I encourage you to follow the links above and read about them).

The degree matrix for $K_n$ is an $n \times n$ diagonal matrix with the value $n-1$ along the diagonal and zeros everywhere else. Write a function to compute the degree matrix for $K_n$ using NumPy.

<!-- grade_id: completegrapha, points: 0, grade: false, solution: true -->
```python
def complete_deg(n):
    """Return the integer valued degree matrix D for the complete graph K_n."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: completegraphb, points: 3, grade: true, solution: false -->
```python
D = complete_deg(5)
assert D.shape==(5,5)
assert D.dtype==np.dtype(int)
assert np.all(D.diagonal()==4*np.ones(5))
assert np.all(D-np.diag(D.diagonal())==np.zeros((5,5),dtype=int))
```

The adjacency matrix for $K_n$ is an $n \times n$ matrix with zeros along the diagonal and ones everywhere else. Write a function to compute the adjacency matrix for $K_n$ using NumPy.

<!-- grade_id: completegraphc, points: 0, grade: false, solution: true -->
```python
def complete_adj(n):
    """Return the integer valued adjacency matrix A for the complete graph K_n."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: completegraphd, points: 3, grade: true, solution: false -->
```python
A = complete_adj(5)
assert A.shape==(5,5)
assert A.dtype==np.dtype(int)
assert np.all(A+np.eye(5,dtype=int)==np.ones((5,5),dtype=int))
```

Use NumPy to explore the eigenvalues or *spectrum* of the Laplacian *L* of $K_n$. What patterns do you notice as $n$ changes? Create a *conjecture* about the general Laplace *spectrum* of $K_n$.

```python
complete_deg(3)
```

<!-- grade_id: completegraphe, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

In a few sentences, propose a *conjecture* about the general Laplace *spectrum* of $K_n$:

<!-- grade_id: completegraphf, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE