# Basic Vector Operations

In this notebook you will implement a number of basic vector operations using pure Python.

```python
import math
```

Implement a function that computes the sum of two input vectors `x` and `y`:


<!-- grade_id: basicvectora, points: 0, grade: false, solution: true -->
```python
def vector_sum(x, y):
    """Compute the vector sum of x and y as a Python list."""
    results = []
    i = 0
    for term in x:
        results.append(term+y[i])
        i += 1
    return results
```

```python
vector_sum([0,0,0,0], [2,2,2,2])
vector_sum([1,2], [-1,-2])
```

Output:
```
[0, 0]
```

<!-- grade_id: basicvectorb, points: 1, grade: true, solution: false -->
```python
assert vector_sum([0,0,0,0], [2,2,2,2])==[2,2,2,2]
assert vector_sum([1,2], [-1,-2])==[0,0]
```

Implement a function that computes the difference of two input vectors `x` and `y`:

<!-- grade_id: basicvectorc, points: 0, grade: false, solution: true -->
```python
def vector_diff(x, y):
    """Compute the vector difference of x and y as a Python list."""
    results = []
    i = 0
    for term in x:
        results.append(term-y[i])
        i += 1
    return results
```

<!-- grade_id: basicvectord, points: 1, grade: true, solution: false -->
```python
assert vector_diff([1,1], [0,0])==[1,1]
assert vector_diff([1,2,3,4],[1,2,3,4])==[0,0,0,0]
```

Implement a function that computes the product of a scalar `a` with a vectors `x`:

<!-- grade_id: basicvectore, points: 0, grade: false, solution: true -->
```python
def vector_scalar_mul(a, x):
    """Compute a*x as a Python list."""
    results = []
    i = 0
    for term in x:
        results.append(term*a)
        i += 1
    return results
```

<!-- grade_id: basicvectorf, points: 1, grade: true, solution: false -->
```python
assert vector_scalar_mul(1.0, [1,2,3,4])==[1,2,3,4]
assert vector_scalar_mul(2.0, [-1,0,1])==[-2,0,2]
```

Implement a function that computes the dot product between two vectors `x` and `y`:

<!-- grade_id: basicvectorg, points: 0, grade: false, solution: true -->
```python
def vector_dot(x, y):
    """Compute the dot product of vectors a and b."""
    results = 0
    i = 0
    for term in x:
        results += term*y[i]
        i += 1
    return results
```

```python
assert vector_dot([0,1,1,2],[1,1,0,2]) == 0+1+0+4
```

<!-- grade_id: basicvectorh, points: 1, grade: true, solution: false -->
```python
assert vector_dot([0,1],[1,0])==0.0
assert vector_dot([1,0],[1,0])==1.0
assert vector_dot([1,1],[-1,1])==0.0
```

Implement a function that computes the $L_1$ norm of a vector `x`:

<!-- grade_id: basicvectori, points: 0, grade: false, solution: true -->
```python
def L1(x):
    """Compute the L1 norm of x."""
    result = 0
    i = 0
    for term in x:
        result += abs(term)
    return result
```

<!-- grade_id: basicvectorj, points: 1, grade: true, solution: false -->
```python
assert L1([0,0])==0
assert L1([1,1])==2
```

Implement a function that computes the $L_2$ norm of a vector `x`:

<!-- grade_id: basicvectork, points: 0, grade: false, solution: true -->
```python
def L2(x):
    """Compute the L2 norm of x."""
    result = 0
    i = 0
    for term in x:
        result += (abs(term)**2)
        result = math.sqrt(result)
    return result
```

<!-- grade_id: basicvectorl, points: 1, grade: true, solution: false -->
```python
assert L2([0,0])==0
assert L2([1,1])==math.sqrt(2)
```

Implement a function that computes the Mean Squared Error (MSE) between two vectors `x` and `y`:

```python
mylist = [1,2,3,4,5]
print (len(mylist)-1)
```

Output:
```
4
```

<!-- grade_id: basicvectorm, points: 0, grade: false, solution: true -->
```python
def mse(x, y):
    """Compute the mean squared error between x and y."""
    result = 0
    i = 0
    for term in x:
        result += (term-y[i])**2
        print (result)
        i += 1
    result /= float(len(mylist))
    return result
```

```python
print (mse([0,0],[1,1]))
```

Output:
```
1
2
0.4
```

<!-- grade_id: basicvectorn, points: 1, grade: true, solution: false -->
```python
assert mse([0,0],[1,1])==1.0
assert mse([1,2,3,4],[5,6,7,8])==16.0
```

Output:
```
1
2
AssertionError                            Traceback (most recent call last)
<ipython-input-19-6ed25a6d37aa> in <module>()
----> 1 assert mse([0,0],[1,1])==1.0
      2 assert mse([1,2,3,4],[5,6,7,8])==16.0

AssertionError:
```

Implement a function that computes the mean of all the elements of a vector `x`:

<!-- grade_id: basicvectoro, points: 0, grade: false, solution: true -->
```python
def mean(x):
    """Compute the mean of the elements of x."""
    for terms in x:
        result += x
    result /= len(x)
    return result
```

<!-- grade_id: basicvectorp, points: 1, grade: true, solution: false -->
```python
assert mean([1,1,1,1])==1
assert mean([1,2,3,4])==2.5
```

Implement a function that computes the variance of the elements of a vector `x`:

<!-- grade_id: basicvectorq, points: 0, grade: false, solution: true -->
```python
def var(x):
    """Compute the variance of the elements of x."""
    for terms in x:
        result += (x - mean(x))**2
    result /= len(x)
    return result
```

<!-- grade_id: basicvectorr, points: 1, grade: true, solution: false -->
```python
assert var([1,1,1,1])==0.0
assert var([1,2,3,4])==1.25
```

Implement a function that computes the standard deviation of a vector `x`:

<!-- grade_id: basicvectors, points: 0, grade: false, solution: true -->
```python
def std(x):
    """Compute the standard deviation of the elements of x."""
    return math.sqrt(var(x))
```

<!-- grade_id: basicvectort, points: 1, grade: true, solution: false -->
```python
assert abs(std([2,4,6,8])-2.23606)<0.0001
```