# Basic Pandas Homework

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
```

## DatFrame and Series

Create a `DataFrame` named `df1` with the following properties:

* 10 rows.
* An `age` column with random ages between 0 and 100 (inclusive).
* A `cell_phone` column of randomly sampled categorical values `ios`, `android`, `windows`.
* A `gender` column of randomly sampled categorical values `f` and `m`.
* The order of columns should be `gender`, `age`, `cell_phone`.
* A row index consisting of lowercase alphabetical letters.

Use the function defined above to create this `DataFrame`.

<!-- grade_id: basic-pandas-a, points: 0, grade: false, solution: true -->
```python
age = np.random.choice(101,10)
```

```python
cell_phone = np.random.choice(('ios','android','windows'),10)
```

```python
gender = np.random.choice(('f','m'), 10)
```

```python

```

```python

```

Make sure your code passes the following `assert` statements:

<!-- grade_id: basic-pandas-b, points: 3, grade: true, solution: false -->
```python
assert 'a' in df1.index
assert 'j' in df1.index
assert 'age' in df1.columns
assert 'cell_phone' in df1.columns
assert 'gender' in df1.columns
assert df1.age.dtype.name=='int64'
assert df1.cell_phone.dtype.name=='category'
assert df1.gender.dtype.name=='category'
assert list(df1.columns)==['gender','age','cell_phone']
```

Create a new `DataFrame`, named `df2` with the following transformations:

* Extract rows `a` through `g`
* Extract the `age` and `gender` column, but put `age` first.
* Reverse the rows to they run `g` to `a`.
* Add a new column named `income` that is a random list of dollars between [0,10000]
* Add a new column named `expenses` that is a random list of dollars between [0,10000]
* Create a new column named `profit` that is computed using the `income` and `expenses` columns.

<!-- grade_id: basic-pandas-c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
df2
```

Make sure your code passes the following `assert` statements:

<!-- grade_id: basic-pandas-d, points: 3, grade: true, solution: false -->
```python
assert df2.index[0]=='g'
assert df2.index[-1]=='a'
assert list(df2.columns)==['age','gender','income','expenses','profit']
assert df2.income.dtype.name=='int64'
assert df2.expenses.dtype.name=='int64'
assert df2.expenses.dtype.name=='int64'
assert all(df2.profit+df2.expenses-df2.income==0)
```

Write a function named `mean_and_std` that takes a `Series` object and returns a new `Series` that contains the mean and standard deviation of the original series.

<!-- grade_id: basic-pandas-e, points: 0, grade: false, solution: true -->
```python
def mean_and_std(s):
    """Compute the mean and std of series s.
    
    Parameters
    ----------
    s : Series
    
    Returns
    -------
    Series containing the mean and std of s with an index of `mean` and `std`.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

Make sure your code passes the following `assert` statements:

<!-- grade_id: basic-pandas-f, points: 2, grade: true, solution: false -->
```python
assert list(mean_and_std(df2.age).index)==['mean','std']
assert mean_and_std(df2.age)['mean']==df2.age.describe()['mean']
assert mean_and_std(df2.age)['std']==df2.age.describe()['std']
```

Use the `.apply()` method with `mean_and_std` to compute the mean and standard deviation of only the columns `income`, `expenses` and profit. Save the result in a variable named `stats`.

<!-- grade_id: basic-pandas-g, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
stats
```

Make sure your code passes the following `assert` statements:

<!-- grade_id: basic-pandas-h, points: 2, grade: true, solution: false -->
```python
d = df2[['income','expenses','profit']].describe()
assert 'income' in stats.columns
assert 'expenses' in stats.columns
assert 'profit' in stats.columns
assert stats.ix['mean','income']==d.ix['mean','income']
assert stats.ix['mean','expenses']==d.ix['mean','expenses']
assert stats.ix['mean','profit']==d.ix['mean','profit']
assert stats.ix['std','income']==d.ix['std','income']
assert stats.ix['std','expenses']==d.ix['std','expenses']
assert stats.ix['std','profit']==d.ix['std','profit']
```